<a href="https://colab.research.google.com/github/LMU-CMSI-1010/ps3-lilscolari/blob/main/Copy_of_ps3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Problem Set 3
#### CMSI 1010: Computer Programming and Lab

# Important Dates

 - Released: 2022-09-29 at 12:00 am [PT]
 - Deadline: 2022-10-06 at 11:59 pm [PT]

# Question 1: Reading Notes

All readings are from the [_Think Python_ textbook](http://greenteapress.com/thinkpython2/thinkpython2.pdf):
* Supporting content for this pset: Chapters 7–9
* Upcoming content: Chapters 11–13, 15–17

Take notes from your readings and submit them as part of this pset.

Chapter 7

•	Variables can be reassigned

•	Python evaluates right side of ‘=’ before assigning it to the variable on the left side.

•	Before you can update a variable, you must initialize it.

•	Increment: adding 1

•	Decrement: subtracting 1

•	Repetition is called iteration

•	Pick up at 7.3

•	Infinite loop results in important value not changing

•	Break jumps out of a loop

•	Algorithm: mechanical process for solving a category of problems

•	Debugging by bisection means splitting code in parts and making debugging 
easier.

Chapter 8

•	Strings are sequence of characters

•	Access a character one at a time with []

•	Expression in brackets is index

•	String index starts at 0

•	Len() returns number of characters in a string

•	Traversal: processing a string one character at a time

•	Segment of string is called a slice

•	Includes number before ‘:’ and but excludes number after ‘:’

•	Omit first index means slice starts at beginning of string

•	Omit second index means slice goes to end of string

•	Strings are immutable (cannot change existing string)

•	String methods take an argument and return a value

•	A method call is an invocation

•	In function is Boolean operator that takes two strings and returns True if the first one appears as a substring in the second

•	Uppercase letters come before lowercase letters in terms of sorting for Python

Chapter 9

•	Open function in Python takes name of a file as parameter and returns file 
object you can use to read the file 

•	Readline method reads characters from a file until a new line is reached and prints the result as a string

•	The empty string is an example of a special case

•	Special case: non-obvious cases where errors can occur

•	Reduction to a previously solved problem: problem you are working on is instance of previously solved problem so application of previous solution works

# Question 2: Understanding loops

### Part A
 For these fragments of code, write the output **_without_** running them.

```python
    num1 = 10
    while num1 > 3:
        print(num1)
        num1 = num1 - 1
```

 #### Output:
    ```
    10
    9
    8
    7
    6
    5
    4
    ```

```python
    num2 = 10
    while True:
        if num2 < 7:
            # break makes us quit the loop
            break
        print(num2)
        num2 = num2 - 1
```

 #### Output:
    ```
    10
    9
    8
    7
    ```

```python
    divisor = 2
    for i in range(0, 10, 2):
        print(i / divisor)
```

 #### Output:
    ```
    0.0
    1.0
    2.0
    3.0
    4.0
    ```

### Part B

These programs use `for` loops to iterate through strings.

```python
    my_string = 'stressed'
    for letter in my_string:
        print(letter)
```

#### Output:
    ```
    s
    t
    r
    e
    s
    s
    e
    d
    ```

```python
    my_string = 'stressed'
    res = ''
    for letter in my_string:
        res = letter + res
    print(res)
```

#### Output:
    ```
    desserts
    ```

Describe, in natural language, what’s happening in this loop:



```
# This is formatted as code
```

This loop iterates through the string and reverses it and assigns the reversed word into a new value.

# Question 3 : Exponentials

Imagine that the `**` operator for numbers has now been removed from the Python programming language. Write a program that prompts the user for a base and an exponent and uses a `for` loop to calculate `base ** exp`. _**Note:** You should handle the case where the user does not type a positive integer._

```python
def exponential(base, exp):

    result = base

    if exp == 0:
        result = 1
    elif exp < 0:
        for i in range(0, exp - 1, -1):
            result = 1 / (result * base)
    else:
        for i in range(0, exp - 1, 1):
            result *= base

    return result
```

```python
base = eval(input('Enter a random number for the base: '))
exp = eval(input('Enter a random number for the exponent: '))
exponential(base, exp)
```

Output:
```
Enter a random number for the base: 5
Enter a random number for the exponent: 3
125
```

The cells below are autograder cells; don't edit those! But you can run them (using the Run button) to test your own code! If they don't spit out any error messages, then your code works!

```python
assert exponential(2,4) == 2**4
```

```python
assert exponential(3,0) == 3**0
```

```python
assert exponential(-2,3) == (-2)**3
```

```python
assert exponential(2,-2) == 2**(-2)
```

```python
import random

num1 = random.randint(0,1000)
num2 = random.randint(0,100)

assert exponential(num1, num2) == num1**num2
```

# Question 4 : Functions and more functions

### Part A

Write a function:

 - named `list_of_positive_indices`
 - takes 1 argument, a list
 - returns 1 value, a list
 
In this function, you should loop through each element of the argument list. If the element is *positive*, then you should append the *index number of the positive element* to a new list. If not, skip that index (i.e. do nothing). Once the loop is finished, return the new list.

For example, `list_of_positive_indices([1, -1, 0, 3])` should return `[0, 3]`. No imports!

*HINT*: Instead of doing a simple `for element in list:` loop, try using the `range()` function in the loop header.

```python
def list_of_positive_indices(in_list):
    list3 = []


    for index in range(0, len(in_list)-1, 1):
        if in_list[index] > 0:
          list3.append(index)
        else:
          pass
    return list3

```

Run the autograder cells below to test your code!

<!-- grade_id: q1c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(986)

list1 = np.random.randint(-100, 100, 10)
indices1 = np.arange(list1.shape[0])[list1 > 0].tolist()
pred1 = list_of_positive_indices(list1.tolist())
assert set(pred1) == set(indices1)
```

<!-- grade_id: q1c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(578752)

list2 = np.random.randint(-100, 100, 100)
indices2 = np.arange(list2.shape[0])[list2 > 0].tolist()
pred2 = list_of_positive_indices(list2.tolist())
assert set(pred2) == set(indices2)
```

### Part B

Write a function that tests if a number is even.

*HINT*: Recall the modulo operator `%`. This performs division on two numbers, but instead of returning their quotient, it returns the *remainder* of the division. For example, `7 % 3` returns 1, since the quotient is 2 with a remainder of 1.

Your function should:

 - be named `is_even`
 - take 1 argument, a number
 - return `True` if the argument is even, `False` otherwise
 
Make sure you're returning a *boolean*!

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def is_even(num):
  if num % 2 == 0:
    return True
  else:
    return False
  
```

<!-- grade_id: q2a_test1, points: 2.5, grade: true, solution: false -->
```python
assert is_even(3) == False
```

<!-- grade_id: q2a_test2, points: 2.5, grade: true, solution: false -->
```python
assert is_even(123456789) == False
```

<!-- grade_id: q2a_test3, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42) == True
```

<!-- grade_id: q2a_test4, points: 2.5, grade: true, solution: false -->
```python
assert is_even(42.1) == False
```

### Part C

How could you test if a number is *odd*, using your `is_even` function from Part B?

Write your answer, either using 1 line of code or explaining in only 1 sentence, in the box below (double click on the statement below to activate the box).

<!-- grade_id: q2b, points: 5, grade: true, solution: true -->


```
If the is_even function returns false, then we know that our function is odd.
```



### Part D *(Optional)*

*This part is optional*. Write a function that finds the *second-smallest* number in a list of numbers. Your function should:

 - be named `second_smallest`
 - take 1 argument, a list of numbers
 - return 1 number, the second-smallest number in the list
 
Take note of that last point. There's a possibility, after all, that the list your function receives as input will have 1 or even 0 elements in it. How, then, do you return the second-smallest element of a 0-element list? Well, simply, you can't. So instead, you'll return False.

For example, `second_smallest([1, 2, 3])` should return `2`. Another example `second_smallest([10, 100, 1000, 10000])` should return 100. And `second_smallest([1])` should returb `False`.

You **cannot** use any built-in sorting functions! Only `range()` and `len()`.

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python
def second_smallest(in_list):
  pass
```

<!-- grade_id: q2c_test1, points: 5, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(89547)

l1 = np.random.randint(-100, 100, 100)
a1 = np.sort(l1)[1]
np.testing.assert_allclose(a1, second_smallest(l1.tolist()))
```

<!-- grade_id: q2c_test2, points: 5, grade: true, solution: false -->
```python
np.random.seed(485)

l2 = np.random.randint(0, 1000, 100)
a2 = np.sort(l2)[1]
np.testing.assert_allclose(a2, second_smallest(l2.tolist()))
```

<!-- grade_id: q2c_test3, points: 2.5, grade: true, solution: false -->
```python
assert second_smallest([]) == False
```

<!-- grade_id: q2c_test4, points: 2.5, grade: true, solution: false -->
```python
assert second_smallest([100]) == False
```

# Question 5: Greetings

### Part A

Print a greeting saying “Hello, Harry” and “Goodbye, Potter” using only the variables provided below. Complete instructions below.

```python
# Write a series of print statements to recreate a 
# greeting box and a goodbye box for Harry Potter.

# *******************
# *                 *
# * Hello, Harry!!! *
# *                 *
# *******************
#
# ***********************
# *                     *
# * Goodbye, Potter!!!  *
# *                     *
# ***********************


name = 'Harry Potter'
greeting = 'Hello, world!'
goodbye = 'Goodbye, all!'
space = ' '
star = '*'


# You may only use the variables defined above. 
# For example, if you wanted to print
# a "T" in stars, you would do:

# print (star * 3)
# print (space + star + space)
# print (space + star + space)

# The answer below would not be acceptable:
# print ('***')
# print (' * ')
# print (' * ')

print(star * 19)
print(star + space * 17 + star)
print(star + space + greeting[0:7] + name[0:5] + greeting[-1] * 3 + space + star)
print(star + space * 17 + star)
print(star * 19)
print()
print(star * 23)
print(star + space * 21 + star)
print(star + space + goodbye[0:9] + name[6:12] + greeting[-1] * 3 + space * 2 + star)
print(star + space * 21 + star)
print(star * 23)
```

Output:
```
*******************
*                 *
* Hello, Harry!!! *
*                 *
*******************

***********************
*                     *
* Goodbye, Potter!!!  *
*                     *
***********************
```

### Part B
 
 **Reflection:** The output of `greetings` looks very similar to the one from PS0. However, the code looks different. How would you compare the two versions? What advantages/disadvantages do you see in the two approaches?

The two versions ultimately result in printing the imagined output but the one in ps3 is much more beneficial. The output in ps3 is constructed primarily from predefined strings. This means that the output can have more variability and can say more phrases. For example, I was able to slice the predefined strings and combine them together to get a different output. The only advantage I see with using the code from ps0 is that you can see exactly what it is you are printing before actually running the program.

# Question 6: Acronym

This exercise will allow you to practice with the methods for strings and lists. An acronym is a word formed by taking the first letters of the words in a phrase and making a word from them. For example, RAM is an acronym for “random access memory.” In this exercise, you are going to write a program that should allow the user to type in a phrase, and the computer should then output the acronym for that phrase. An example of how your program should work is as follows:

```
    Enter a phrase: loyola marymount university
    Your acronym is: LMU
```

**Note:** The acronym should be all uppercase, even if the words in the phrase are not capitalized.

```python
phrase = input('Enter a phrase: ').upper()
list = phrase.split()

acronym = ''

for word in list:
    acronym = acronym + word[0:1]

print(acronym)
```

Output:
```
Enter a phrase: i love loyola marymount university
ILLMU
```